# This is the main input for unit testing.
# Only this file is passed as command line argument.

rule_files:
  - ../rules/alerts.yml

evaluation_interval: 1m

tests:
  # Test 1.
  - interval: 1m
    # Series data.
    input_series:
      - series: 'redis_connected_clients{job="redis", instance="localhost:6379"}'
        values: '420 480 520 540 560 580 600 520 480 520 460 440 380 400 500'

    # Unit test for alerting rules.
    alert_rule_test:
      # Unit test 1.
      - eval_time: 7m
        alertname: ShardConnectionCount
        exp_alerts:
          # Alert 1.
          - exp_labels:
              severity: warning
              instance: localhost:6379
              job: redis
            exp_annotations:
              summary: "Shards - Excessive Connection Count"
              description: Shard localhost:6379 exceeded 500 connections for 5 minutes
