# This is the main input for unit testing.
# Only this file is passed as command line argument.

rule_files:
  - ../rules/alerts.yml

evaluation_interval: 1m

tests:
  # Test 1.
  - interval: 1m
    # Series data.
    input_series:
      - series: 'bdb_crdt_syncer_status{job="redis", instance="localhost:6379"}'
        values: '0 1 2 2 1 2 1 1 2 2 2 1 2 0 1'

    # Unit test for alerting rules.
    alert_rule_test:
      # Unit test 1.
      - eval_time: 10m
        alertname: CRDTSync
        exp_alerts:
          # Alert 1.
          - exp_labels:
              severity: warning
              instance: localhost:6379
              job: redis
            exp_annotations:
              summary: Replication - Unsynchronized (CRDT)
              description: CRDT Replication on localhost:6379 not synchronized for 5 minutes
