# This is the main input for unit testing.
# Only this file is passed as command line argument.

rule_files:
  - ../rules/alerts.yml

evaluation_interval: 1m

tests:
  # Test 1.
  - interval: 5m
    # Series data.
    input_series:
      - series: 'bdb_crdt_syncer_local_ingress_lag_time{job="redis", cluster="localhost"}'
        values: '420 580 520 540 560 580 600 520 510 520 530 550 580 400 500'

    # Unit test for alerting rules.
    alert_rule_test:
      # Unit test 1.
      - eval_time: 11m
        alertname: CRDTLag Status
        exp_alerts:
          # Alert 1.
          - exp_labels:
              severity: warning
              job: redis
              cluster: localhost
            exp_annotations:
              summary: Replication - High Latency (CRDT)
              description: CRDT Replication Latency on localhost exceeded 500ms for 10 minutes
