# This is the main input for unit testing.
# Only this file is passed as command line argument.

rule_files:
  - ../rules/alerts.yml

evaluation_interval: 1m

tests:
  # Test 1.
  - interval: 1m
    # Series data.
    input_series:
      - series: 'redis_master_link_status{job="redis", role="slave", cluster="localhost"}'
        values: '0 0 0 0 0 0 0 0 0 0 0 0 0 0 0'

    # Unit test for alerting rules.
    alert_rule_test:
      # Unit test 1.
      - eval_time: 1m
        alertname: Master Shard Down
        exp_alerts:
          # Alert 1.
          - exp_labels:
              severity: notification
              job: redis            
              cluster: localhost
              role: slave            
            exp_annotations:
              summary: Master Shard Down
              description: "Slave has no master - Cluster: {{ $labels.cluster }} Shard: {{$labels.redis}} Node: {{$labels.node}}"
