# This is the main input for unit testing.
# Only this file is passed as command line argument.

rule_files:
  - ../rules/alerts.yml

evaluation_interval: 1m

tests:
  # Test 1.
  - interval: 1m
    # Series data.
    input_series:
      - series: 'bdb_conns{job="redis", cluster="localhost", bdb="1"}'
        values: '0 0 0 0 0 0 0 0 0 0 0 0 0 0 0'


    # Unit test for alerting rules.
    alert_rule_test:
      # Unit test 1.
      - eval_time: 1m
        alertname: No Redis Connections
        exp_alerts:
          # Alert 1.
          - exp_labels:
              severity: notification
              job: redis            
              cluster: localhost
              bdb: 1            
            exp_annotations:
              summary: "No Redis Connections"
              description: "No Connections - Cluster: localhost DB: 1 Connections: 0"
              runbook: https://github.com/Redislabs-Solution-Architects/RedisEnterprisePrometheus/blob/master/runbooks/connections.md#no-redis-connections
