# This is the main input for unit testing.
# Only this file is passed as command line argument.

rule_files:
  - ../rules/alerts.yml

evaluation_interval: 1m

tests:
  # Test 1.
  - interval: 1m
    # Series data.
    input_series:
      - series: 'bdb_total_req{job="redis", instance="localhost:6379"}'
        values: '0 0 0 0 0 0 0 0 0 0 0 0 0 0 0'

    # Unit test for alerting rules.
    alert_rule_test:
      # Unit test 1.
      - eval_time: 1m
        alertname: No Redis Requests
        exp_alerts:
          # Alert 1.
          - exp_labels:
              severity: critical
              instance: localhost:6379
              job: redis            
            exp_annotations:
              summary: No Redis Requests
              description: "Too few Redis operations - Cluster:  DB:   0 (ops/sec)"
              runbook: https://github.com/Redislabs-Solution-Architects/RedisEnterprisePrometheus/blob/master/runbooks/throughput.md#no-redis-requests
