# This is the main input for unit testing.
# Only this file is passed as command line argument.

rule_files:
  - ../rules/alerts.yml

evaluation_interval: 1m

tests:
  # Test 1.
  - interval: 30s
    # Series data.
    input_series:
      - series: 'node_ephemeral_storage_free{job="redis", cluster="localhost"}'
        values: '3 2 1 1 1 1 1 1 1 2 2 2 2 2 3'
      - series: 'node_ephemeral_storage_avail{job="redis", cluster="localhost"}'
        values: '500 500 500 500 500 500 500 500 500 500 500 500 500 500 500'

    # Unit test for alerting rules.
    alert_rule_test:
      # Unit test 1.
      - eval_time: 2m
        alertname: Node Ephemeral Storage
        exp_alerts:
          # Alert 1.
          - exp_labels:
              severity: critical
              job: redis
              cluster: localhost            
            exp_annotations:
              summary: Node Ephemeral Storage
              description: "Low on Ephemeral Storage - Cluster: localhost Space Free: 0%"
              runbook: https://github.com/Redislabs-Solution-Architects/RedisEnterprisePrometheus/blob/master/runbooks/nodes.md#node-ephemeral-storage
