# This is the main input for unit testing.
# Only this file is passed as command line argument.

rule_files:
  - ../rules/alerts.yml

evaluation_interval: 1m

tests:
  # Test 1.
  - interval: 1m
    # Series data.
    input_series:
      - series: 'node_persistent_storage_free{job="redis", instance="localhost:6379"}'
        values: '50 40 30 20 10 5 5 5 5 5 5 5 5 10 20'
      - series: 'node_persistent_storage_avail{job="redis", instance="localhost:6379"}'
        values: '50 40 30 20 10 5 5 5 5 5 5 5 5 10 20'

    # Unit test for alerting rules.
    alert_rule_test:
      # Unit test 1.
      - eval_time: 2m
        alertname: Node Persistent Storage
        exp_alerts:
          # Alert 1.
          - exp_labels:
              severity: critical
              instance: localhost:6379
              job: redis            
            exp_annotations:
              summary: Node Persistent Storage
              description: "Low on Persistent Storage - Cluster: {{ $labels.cluster }} Space Free: {{$value}}%"
