# This is the main input for unit testing.
# Only this file is passed as command line argument.

rule_files:
  - ../rules/alerts.yml

evaluation_interval: 1m

tests:
  # Test 1.
  - interval: 1m
    # Series data.
    input_series:
      - series: 'bdb_replicaof_syncer_local_ingress_lag_time{job="redis", instance="localhost:6379"}'
        values: '420 580 520 540 560 580 600 520 510 520 530 550 580 400 500'

    # Unit test for alerting rules.
    alert_rule_test:
      # Unit test 1.
      - eval_time: 12m
        alertname: ReplicaLag
        exp_alerts:
          # Alert 1.
          - exp_labels:
              severity: warning
              instance: localhost:6379
              job: redis
            exp_annotations:
              summary: Replication - High Latency (Status)
              description: Replication Latency localhost:6379 exceeded 500ms for 10 minutes
