# This is the main input for unit testing.
# Only this file is passed as command line argument.

rule_files:
  - ../rules/alerts.yml

evaluation_interval: 1m

tests:
  # Test 1.
  - interval: 1m
    # Series data.
    input_series:
      - series: 'redis_up{job="redis", instance="localhost:6379"}'
        values: '1 1 0 0 0 0 0 0 0 0 0 0 0 1 1'

    # Unit test for alerting rules.
    alert_rule_test:
      # Unit test 1.
      - eval_time: 1m
        alertname: InstanceDown
        exp_alerts:
          # Alert 1.
          - exp_labels:
              severity: page
              instance: localhost:6379
              job: redis            
            exp_annotations:
              summary: Redis Shard instance is down
              description: Redis Shard instance {{ $labels.instance }} is down
