# This is the main input for unit testing.
# Only this file is passed as command line argument.

rule_files:
  - ../rules/alerts.yml

evaluation_interval: 30s

tests:
  # Test 1.
  - interval: 30s
    # Series data.
    input_series:
      - series: 'bdb_evicted_objects{job="redis", cluster="localhost", bdb="1"}'
        values: '1 1 2 3 4 5 6 7 6 5 4 3 2 1 1 2 4 3 5 2 5 7 8 9'

    # Unit test for alerting rules.
    alert_rule_test:
      # Unit test 1.
      - eval_time: 4m
        alertname: Unexpected Object Eviction
        exp_alerts:
          # Alert 1.
          - exp_labels:
              severity: critical
              job: redis            
              cluster: localhost
              bdb: 1
            exp_annotations:
              summary: Unexpected Object Eviction
              description: "Evictions Occuring: localhost DB: 1 EvictionsPerSecond: 6"
              runbook: https://github.com/Redislabs-Solution-Architects/RedisEnterprisePrometheus/blob/master/runbooks/utilization.md#objects-being-evicted-from-database